Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, email and UFID.
Please do not modify instruction cells or any cells with automated tests (marked with `[ASSERTS]`). Note: you can add new cells if you need them, but answers must be in the cells with `YOUR CODE HERE` or "YOUR ANSWER HERE" comments.

---

## Homework 2: Regression and Ensembles

## Preamble: Write your Name, Email and UFID

<!-- grade_id: name-email-ufid, points: 0, grade: false, solution: true -->
```python
NAME = 'Tre\' Jeter'
EMAIL = 't.jeter@ufl.edu'
UFID = 19469876

# YOUR CODE HERE
# raise NotImplementedError()

print('Homework 2 -- name: {}, email: {}, UFID: {}\n'.format(NAME, EMAIL, UFID))
```

Output:
```
Homework 2 -- name: Tre' Jeter, email: t.jeter@ufl.edu, UFID: 19469876
```

<!-- grade_id: name-email-ufid-asserts, points: 0, grade: true, solution: false -->
```python
""" [ASSERTS] Check that your name, email, and UFID is filled in."""
assert NAME != '' and NAME != 'Your name here.' and len(NAME) > 3
assert EMAIL != '' and EMAIL != 'Your email here.' and len(EMAIL) > 7
assert type(UFID) == int and UFID != 12345678 and UFID >= 10000000 and UFID <= 99999999
```

## Academic Integrity

### <span style="color:red;">This is an individual assignment. Academic integrity violations (i.e., cheating, plagiarism) will be reported to SCCR!</span><br/>
#### The official CISE policy recommended for such offenses is a course grade of E. Additional sanctions may be imposed by SCCR such as marks on your permanent educational transcripts, dismissal or expulsion.
#### Reminder of the Honor Pledge: On all work submitted for credit by Students at the University of Florida, the following pledge is either required or implied: *"On my honor, I have neither given nor received unauthorized aid in doing this assignment."*

#### Acknowledgement: Do you acknowledge and understand the academic integrity warning above? 

<!-- grade_id: academic-integrity, points: 0, grade: false, solution: true -->
```python
academic_integrity_acknowledgement = True
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: academic-integrity-assert, points: 0, grade: true, solution: false -->
```python
""" [ASSERTS] Check that you acknowledge the academic integrity warning, you understand it and have been reminded of the UF Honor Pledge."""
assert academic_integrity_acknowledgement
```

---
# [Task 1] (20 points) Loading and preprocessing.

### [Task 1] We will use the Bike Sharing dataset (hourly). A version of this dataset is included in the homework handout archive.
### This dataset contains features of users bike sharing/rental on an hourly basis.
### The task is to predict how many users are sharing/renting a bike.
### In this task you will load the data and preprocess it.

#### The following cell's code (import statements etc.) is provided for you and you should not need to change it.

```python
# Load packages we need
import sys
import os
import time

import numpy as np
import pandas as pd
import sklearn

from matplotlib import pyplot as plt
plt.rcParams.update({'font.size': 16})

# Let's check our software versions
print('------------')
print('### Python version: ' + __import__('sys').version)
print('### NumPy version: ' + np.__version__)
print('### Scikit-learn version: ' + sklearn.__version__)
print('------------')

def var_exists(var_name):
    return (var_name in globals() or var_name in locals())
```

Output:
```
------------
### Python version: 3.9.16 (main, Jan 11 2023, 16:16:36) [MSC v.1916 64 bit (AMD64)]
### NumPy version: 1.20.3
### Scikit-learn version: 1.4.0
------------
```

#### This is the seed we will use, do not change it.

```python
# set the seed
seed = 42
np.random.seed(seed)

prop_vec = [16, 2, 2] # proportions for train - val - test splits
```

<!-- grade_id: seed_checking, points: 0, grade: true, solution: false -->
```python
""" [ASSERTS] Check seed. """
assert seed == 42
```

### Loading data (set the path correctly so it runs on your machine --- don't submit the data file with your notebook).
#### Note: this dataset has missing values (artificially introduced), which you'll need to fill in before you can train a model

<!-- grade_id: task1-loaddata, points: 0, grade: false, solution: true -->
```python
"""Fill in the path to the directory where 'bikesharehour.csv.gz' is located.
"""
# data_root = '<relative or absolute path to directory>' #put the path here
data_root = 'data/'
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: task1-loaddata-test, points: 0, grade: true, solution: false -->
```python

dataset_fp = os.path.join(data_root, 'bikesharehour.csv.gz')
assert os.path.exists(dataset_fp), 'Dataset not found ({})!'.format(dataset_fp)
df = pd.read_csv(dataset_fp, compression='gzip', header=0, na_values='?')

# Check that we loaded the data as expected
df_expected_shape = (17379, 15)
assert df.shape == df_expected_shape, 'Unexpected shape of df!'

df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 17379 entries, 0 to 17378
Data columns (total 15 columns):
 #   Column      Non-Null Count  Dtype  
---  ------      --------------  -----  
 0   season      16320 non-null  float64
 1   year        16231 non-null  float64
 2   month       16304 non-null  float64
 3   hour        16254 non-null  float64
 4   holiday     16277 non-null  float64
 5   weekday     16282 non-null  float64
 6   workingday  16297 non-null  float64
 7   weathersit  16324 non-null  float64
 8   temp        16242 non-null  float64
 9   atemp       16271 non-null  float64
 10  hum         16252 non-null  float64
 11  windspeed   16281 non-null  float64
 12  registered  16244 non-null  float64
 13  nsqrtc      16263 non-null  float64
 14  count       17379 non-null  int64  
dtypes: float64(14), int64(1)
memory usage: 2.0 MB
```

## [Task 1a] (5 points) Show the first 8 rows of the dataframe df. 

<!-- grade_id: task1a_answer, points: 0, grade: false, solution: true -->
```python
"""Put your code here.
"""
## what does the data look like?
# YOUR CODE HERE
df.head(8)
# raise NotImplementedError()
```

Output:
```
season  year  month  hour  holiday  weekday  workingday  weathersit  temp  \
0     1.0   0.0    NaN   0.0      0.0      6.0         0.0         1.0   NaN   
1     1.0   0.0    NaN   1.0      0.0      6.0         0.0         1.0   NaN   
2     1.0   0.0    1.0   2.0      0.0      6.0         0.0         1.0   0.0   
3     1.0   0.0    1.0   3.0      0.0      6.0         0.0         1.0   0.0   
4     1.0   0.0    1.0   4.0      0.0      6.0         0.0         1.0   0.0   
5     1.0   0.0    1.0   5.0      0.0      6.0         0.0         2.0   0.0   
6     1.0   0.0    1.0   6.0      0.0      6.0         0.0         1.0   0.0   
7     1.0   0.0    1.0   7.0      0.0      6.0         0.0         1.0   0.0   

   atemp  hum  windspeed  registered  nsqrtc  count  
0    0.0  0.0        0.0        13.0    -5.0     16  
1    0.0  0.0        0.0        32.0    -8.0     40  
2    0.0  0.0        0.0        27.0    -7.0     32  
3    0.0  0.0        0.0        10.0    -5.0     13  
4    0.0  0.
```

<!-- grade_id: task1-features, points: 5, grade: true, solution: false -->
```python
# grab all the data as a numpy matrix
all_xy = df.to_numpy()

col_names = [c for c in df.columns]
features = col_names[:-1]
target = col_names[-1]

# what are the features and what is the target?
print('features: {} --- target: {}'.format(features, target))

```

Output:
```
features: ['season', 'year', 'month', 'hour', 'holiday', 'weekday', 'workingday', 'weathersit', 'temp', 'atemp', 'hum', 'windspeed', 'registered', 'nsqrtc'] --- target: count
```

## [Task 1b] (5 points) Answer the following questions by setting variables (hardcoding or computing the answer) from 'all_xy'. (1) If we do supervise learning to predict 'count' based on our features is this classification or regression? (Set the corresponding variable to True.) (2) How many NaNs are there in columns 'holiday','temp', and 'count'?

<!-- grade_id: task1b-code, points: 0, grade: false, solution: true -->
```python
"""Put your code here and set the variables appropriately.
"""
classification = False
regression = True

h_NaNs = df['holiday'].isna().sum()
t_NaNs = df['temp'].isna().sum()
cnt_NaNs = df['count'].isna().sum()

holiday_NaNs = h_NaNs
temp_NaNs = t_NaNs
count_NaNs = cnt_NaNs

# holiday_NaNs, temp_NaNs, count_NaNs

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: task1b_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check 1b completed. """

assert var_exists('classification') and var_exists('regression') and classification != regression
assert var_exists('holiday_NaNs') and holiday_NaNs >= 0
assert var_exists('temp_NaNs') and temp_NaNs >= 0
assert var_exists('count_NaNs') and count_NaNs >= 0

```

### [Task 1c] (5 points) Let's impute the missing values! Use Scikit-learn's SimpleImputer to replace all NaNs in 'all_xy' with the *most frequent* value in each column. Use copy=True and store the results in 'all_xy_noNaNs' 

<!-- grade_id: task1c_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~2-3 lines).
"""

# YOUR CODE HERE
from sklearn.impute import SimpleImputer

# imputer = SimpleImputer(strategy='most_frequent')
imputer = SimpleImputer(copy=True)
all_xy_noNaNs = imputer.fit_transform(all_xy)

# raise NotImplementedError()
```

<!-- grade_id: task1c_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check 1c completed. """

assert var_exists('all_xy_noNaNs') and all_xy_noNaNs.shape == df_expected_shape

```

### [Task 1d] (5 points) Min-max normalize the features (to [0,1] range) using sklearn's MinMaxScaler. Store the results into 'scaled_all_x'. Then split the data into train, val, test according to the proportion in prop_vec using sklearn's train_test split. Store the results into 'train_x', 'train_y', 'test_x', 'test_y', 'val_x', 'val_y'. 

<!-- grade_id: task1d_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~5-7 lines).
"""

from sklearn.preprocessing import MinMaxScaler
from sklearn.model_selection import train_test_split

scaler = MinMaxScaler()
scaled_all_x = scaler.fit_transform(all_xy_noNaNs[:,:-1])

target_value = all_xy_noNaNs[:,-1]

train_x, test_x, train_y, test_y = train_test_split(scaled_all_x, target_value, test_size=(prop_vec[1] + prop_vec[2]) / (prop_vec[0] + prop_vec[1] + prop_vec[2]), random_state=seed)
test_x, val_x, test_y, val_y = train_test_split(test_x, test_y, test_size=(prop_vec[1] + prop_vec[2]) / (prop_vec[0]/prop_vec[1]), random_state=seed)

train_x.shape, train_y.shape, val_x.shape, val_y.shape, test_x.shape, test_y.shape
# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
((13903, 14), (13903,), (1738, 14), (1738,), (1738, 14), (1738,))
```

<!-- grade_id: task1d_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check split for task 1d. """
assert var_exists('scaled_all_x')
assert var_exists('train_x') and var_exists('train_y') and train_x.shape[0] == train_y.shape[0]
assert var_exists('val_x') and var_exists('val_y') and val_x.shape[0] == val_y.shape[0]
assert var_exists('test_x') and var_exists('test_y') and test_x.shape[0] == test_y.shape[0]
assert train_x.shape == (13903, 14) and val_x.shape == (1738, 14) and test_x.shape == (1738, 14)
assert np.amin(train_x) >= 0.0 and np.amax(train_x) <= 1.0

```

## [Task 2] (30 points) Evaluating and training linear models.

### [Task 2a] (5 points) Fill in the code to calculate and return the metrics.

<!-- grade_id: task2a_code, points: 0, grade: false, solution: true -->
```python
"""Fill in your code below (~5-7 lines).
"""
from sklearn.metrics import mean_squared_error, mean_absolute_error, r2_score

def r2_mse_mae_eval(model, tr_x, tr_y, v_x, v_y, pref='', verb=True):

    """Fill in your code below (~1-2 lines).
    """
    # R^2 the coefficient of determination
    # YOUR CODE HERE
    train_r2 = r2_score(tr_y, model.predict(tr_x))
    val_r2 = r2_score(v_y, model.predict(v_x))
    # raise NotImplementedError()
    
    if verb:
        print('{}Train R^2: {:.3f}, Val  R^2: {:.3f}'.format(pref, train_r2, val_r2))

    train_pred = model.predict(tr_x)
    val_pred = model.predict(v_x)

    """Fill in your code below (~1-2 lines).
    """
    # YOUR CODE HERE
    train_mse = mean_squared_error(tr_y, train_pred)
    val_mse = mean_squared_error(v_y, val_pred)
    # raise NotImplementedError()

    if verb:
        print('{}Train MSE: {:.3f}, Val MSE: {:.3f}'.format(pref, train_mse, val_mse))

    """Fill in your code below (~1-2 lines).
    """
    # YOUR CODE HERE
    train_mae = mean_absolute_error(tr_y, train_pred)
    val_mae = mean_absolute_error(v_y, val_pred)
    # raise NotImplementedError()

    if verb:
        print('{}Train MAE: {:.3f}, Val MAE: {:.3f}'.format(pref, train_mae, val_mae))

    return train_r2, val_r2, train_mse, val_mse, train_mae, val_mae
```

<!-- grade_id: task2a_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check split for task 2a. """
assert var_exists('r2_mse_mae_eval')

```

### [Task 2b] (5 points) Train a linear regression model using the default (hyper)parameters. Call the resulting trained model 'lrmodel'.

<!-- grade_id: task2b_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~1-2 lines).
"""
from sklearn.linear_model import LinearRegression, Ridge

# YOUR CODE HERE
lrmodel = LinearRegression().fit(train_x, train_y)
# raise NotImplementedError()

_ = r2_mse_mae_eval(lrmodel, train_x, train_y, val_x, val_y, pref='[{}] '.format(lrmodel.__class__.__name__))
```

Output:
```
[LinearRegression] Train R^2: 0.895, Val  R^2: 0.891
[LinearRegression] Train MSE: 3483.086, Val MSE: 3244.993
[LinearRegression] Train MAE: 34.013, Val MAE: 32.707
```

### [Task 2b] How good is that model? (A few sentences is fine.)

<!-- grade_id: cell-task2b_answer, points: 5, grade: true, solution: true -->
```python
""" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"""
# 
## Answer: 
# The r^2 value on both training and validation sets are both high which means the model
# is generalizing to the target pretty well. However, the MSE and MAE values are extremely
# high meaning that the model's predictive capabilities are not that good. MSE and MAE should
# be smaller values so the fact that the reported values are really high means the accuracy of 
# this model is probably not that good. So, this model isn't that good in that regard.

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"
```

### [Task 2c] (5 points) Fill in the code below to setup a grid search. Concatenate the train and val sets into 'search_x' and 'search_y'.

<!-- grade_id: task2c_code, points: 0, grade: false, solution: true -->
```python
## some code to do a grid search and automatically train & evaluate the model with the best hyperparams.
from sklearn.model_selection import GridSearchCV

def do_grid_search(model, param_grid, x, y):
    gs = GridSearchCV(model, param_grid, scoring='neg_mean_squared_error')
    gs_res = gs.fit(x, y)
    return  gs_res.best_params_


def search_train_eval(model, param_grid, tr_x=train_x, tr_y=train_y, v_x=val_x, v_y=val_y):

    """Put your code here (~2-3 lines). Concatenate the train and val sets into 'search_x' and 'search_y'.
    """
    # YOUR CODE HERE
    search_x = np.concatenate((train_x, val_x))
    search_y = np.concatenate((train_y, val_y))
    # raise NotImplementedError()
    
    hyperparams = do_grid_search(model, param_grid, search_x, search_y)
    
    class_obj = type(model)
    m = class_obj(**hyperparams).fit(tr_x, tr_y)
    
    cn = str(class_obj).split("'")[1]
    cn = cn.split('.')[-1]
    print('{}({})'.format(cn, hyperparams))

    r2_mse_mae_eval(m, tr_x, tr_y, v_x, v_y, pref='\t')

    return m
```

<!-- grade_id: task2c_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 2c. """
assert var_exists('search_train_eval')

```

### [Task 2d] (10 points) Take a look at the code of search_train_eval() and do_grid_search(). Answer the following questions: 
### 1. Can you do LAD regression with sklearn? Why or why not? (Justify your answer.) Hint: take a look at the sklearn documentation and the course slides.
### 2. Why is the scoring function for the grid search 'neg_mean_squared_error' (as opposed to 'mean_squared_error')? 
### 3. Why is it okay to do the search over search_x and search_y which are the concatenation of the train and 'validation sets?

<!-- grade_id: task2d_answer, points: 10, grade: true, solution: true -->
```python
""" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"""
# 
### Hint: take a look at the documentation of scikit-learn and think.
## Answer: 
# (1) From the documentation,it seems that LADRegression cannot be done directly with sklearn 
# (i.e.,from sklearn.linear_model import LADRegression), but there are other ways to get similar results 
# such as using the HuberRegressor or SGDRegressor.
# https://gurobi-optimization-gurobi-optimods.readthedocs-hosted.com/en/latest/mods/lad-regression.html
# https://scikit-lego.netlify.app/linear-models.html#Least-Absolute-Deviation-Regression
# https://stats.stackexchange.com/questions/376643/is-there-any-library-for-least-absolute-deviation-lad-regression-with-regulari#:~:text=If%20you%20use%20the%20SGDRegressor,to%20LAD%20with%20L2%20regularization.&text=QuantileRegressor%20in%20scikit%2Dlearn%20is,heavily%2C%20by%20default)%20regularized.
# From the lecture slides, LADRegression is more robust to outliers, but mathematically less convenient to work with than OLS.

# (2) The usual convention of sklearn is that higher scores means better performance.
# Because MSE is a loss function, the lower the scores, the better the performance. Therefore,
# by negating the MSE, we can ensure the score is truly maximized as sklearn intends while reporting
# an accurate value.

# (3) It's okay to search over search_x and search_y because the validation sets main purpose is for
# hyperparamter tuning and model evaluation. Having it included in the search allows for 
# hyperparameter optimization. 

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"
```

### [Task 2e] (5 points) Write the code below to do a grid search on a LassoLars model. Call the resulting model 'lassolars_model'. Be sure to tune the regularization constant and 'fit_intercept'. Make sure training converges and that you set the seed.

<!-- grade_id: task2e_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~2-3 lines).
"""

# CHECK AGAIN
from sklearn.linear_model import LassoLars

# YOUR CODE HERE
param_grid = {
    'alpha': [0.01, 0.1, 0.5, 1.0, 5.0, 10.0],
    'fit_intercept': [True, False]
}

lassolars_model = LassoLars(random_state=seed)
lassolars_hyperparams = do_grid_search(lassolars_model, param_grid, train_x, train_y)
lassolars_model_best = LassoLars(**lassolars_hyperparams).fit(train_x, train_y)

print("Best Hyperparameters:", lassolars_hyperparams)
search_train_eval(lassolars_model, param_grid, train_x, train_y, val_x, val_y)

def isConverged(model):
    return hasattr(model, 'coef_')

isConverged(lassolars_model_best)
# raise NotImplementedError()
```

Output:
```
Best Hyperparameters: {'alpha': 0.01, 'fit_intercept': True}
LassoLars({'alpha': 0.01, 'fit_intercept': True})
	Train R^2: 0.895, Val  R^2: 0.891
	Train MSE: 3483.127, Val MSE: 3243.863
	Train MAE: 33.991, Val MAE: 32.680
True
```

<!-- grade_id: task2e_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 2e. """
assert var_exists('lassolars_model')

```

## [Task 3] (25 points) Let's train polynomial regression models!

### [Task 3a] (15 points) Use PolynomialFeatures to create a version of the data with all features of degree 3. Follow the provided instructions.

<!-- grade_id: task3a-code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~10-15 lines).
    1. Use PolynomialFeatures to create a version of the data with all features of degree 3. Make sure to allow interactions (interaction_only=False) and set include_bias=False.
    Store the result in 'all_x_pf'. Ensure that you make a copy of the original data and you use the scaled features ('scaled_all_x')!
    2. Split the data ('all_x_pf') into train-val-test using proportion from 'prop_vec' and save the result as 'train_{x,y}_pf', 'val_{x,y}_pf', and 'test_{x,y}_pf'.
    3. Train three models for comparison. 
        (a) The first is a ridge regression model 'ridge' on the original scaled data (no polynomial features) where you tuned the hyperparameters.
        (b) The second is a linear regression model 'lr_pf' on the polynomial features data.
        (c) The third is a regularized "linear" model (ridge regression) on the polynomial features data where you tuned the hyperparameters (including regularization constant ensuring alpha >= 1.0).
"""
# YOUR CODE HERE
from sklearn.preprocessing import PolynomialFeatures
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression, Ridge

poly = PolynomialFeatures(degree=3, interaction_only=False, include_bias=False)
all_x_pf = poly.fit_transform(scaled_all_x)

train_x_pf, test_x_pf, train_y_pf, test_y_pf = train_test_split(all_x_pf, target_value, test_size=(prop_vec[1] + prop_vec[2]) / (prop_vec[0] + prop_vec[1] + prop_vec[2]), random_state=seed)
test_x_pf, val_x_pf, test_y_pf, val_y_pf = train_test_split(test_x_pf, test_y_pf, test_size=(prop_vec[1] + prop_vec[2]) / (prop_vec[0]/prop_vec[1]), random_state=seed)

print("Ridge w/o Polynomial Features:")
ridge_nopf = Ridge().fit(scaled_all_x, target_value)
search_train_eval(ridge_nopf, {}, scaled_all_x, target_value, val_x, val_y)

print("\nLinear Regression w/ Polynomial Features:")
lr_pf = LinearRegression().fit(train_x_pf, train_y_pf)
search_train_eval(lr_pf, {}, train_x_pf, train_y_pf, val_x_pf, val_y_pf)

print("\nRidge w/ Polynomial Features:")
ridge_pf = Ridge(alpha=1.0).fit(train_x_pf, train_y_pf)
ridge_pf_params = {'alpha': [1.0, 3.0, 5.0, 7.0, 10.0]}
ridge_pf_best_params = do_grid_search(ridge_pf, ridge_pf_params, train_x_pf, train_y_pf)
ridge_pf_best_model = Ridge(**ridge_pf_best_params).fit(train_x_pf, train_y_pf)
search_train_eval(ridge_pf, ridge_pf_params, train_x_pf, train_y_pf, val_x_pf, val_y_pf)

# raise NotImplementedError()
```

Output:
```
Ridge w/o Polynomial Features:
Ridge({})
	Train R^2: 0.895, Val  R^2: 0.892
	Train MSE: 3465.565, Val MSE: 3230.597
	Train MAE: 33.753, Val MAE: 32.515

Linear Regression w/ Polynomial Features:
LinearRegression({})
	Train R^2: 0.955, Val  R^2: 0.948
	Train MSE: 1478.207, Val MSE: 1556.764
	Train MAE: 17.189, Val MAE: 17.937

Ridge w/ Polynomial Features:
Ridge({'alpha': 1.0})
	Train R^2: 0.951, Val  R^2: 0.945
	Train MSE: 1627.915, Val MSE: 1640.276
	Train MAE: 16.988, Val MAE: 16.902
Ridge()
```

<!-- grade_id: task3a-tests, points: 15, grade: true, solution: false -->
```python
""" [ASSERTS] Checks for task 3a. """
assert var_exists('all_x_pf') and all_x_pf.shape == (17379, 679)
assert var_exists('train_x_pf') and var_exists('train_y_pf') and train_x_pf.shape[0] == train_y_pf.shape[0]
assert var_exists('val_x_pf') and var_exists('val_y_pf') and val_x_pf.shape[0] == val_y_pf.shape[0]
assert var_exists('test_x_pf') and var_exists('test_y_pf') and test_x_pf.shape[0] == test_y_pf.shape[0]
assert train_x_pf.shape[0] == 13903 and val_x_pf.shape[0] == 1738 and test_x_pf.shape == val_x_pf.shape
assert train_x_pf.shape[1] == val_x_pf.shape[1]
assert np.amin(train_x_pf) >= 0.0 and np.amax(train_x_pf) <= 1.0

assert var_exists('ridge_nopf') and var_exists('lr_pf') and var_exists('ridge_pf')
assert ridge_nopf.coef_.shape == (14,)

```

### [Task 3b] (5 points) For each of the three models, print the three most important features (coef values and name). You can use 'get_feature_names_out'.

<!-- grade_id: task3b_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~5-7 lines).
"""
# YOUR CODE HERE
feature_names_pf = poly.get_feature_names_out(input_features=col_names[:-1])

models = [ridge_nopf, lr_pf, ridge_pf_best_model]
model_names = ['Ridge Regression Model on Original Scaled Data:',
               'Linear Regression Model on Polynomial Features Data:',
               'Regularized Linear Model on Polynomial Features Data:']

for model, name in zip(models, model_names):
    print(name)
    sorted_indices = np.argsort(np.abs(model.coef_))[::-1][:3]
    for idx in sorted_indices:
        print(f"{feature_names_pf[idx]}: {model.coef_[idx]}")
    print()

# raise NotImplementedError()
```

Output:
```
Ridge Regression Model on Original Scaled Data:
registered: 1000.6557404786477
hour: 37.50448010500321
workingday: -35.413610212240044

Linear Regression Model on Polynomial Features Data:
workingday^2: 393893214467.2939
workingday^3: -234246012137.10712
workingday: -159647202284.1828

Regularized Linear Model on Polynomial Features Data:
nsqrtc: -577.3443035350554
registered: 544.5841731160823
nsqrtc^3: 463.6812378197391
```

<!-- grade_id: task3b_manual_answer, points: 5, grade: true, solution: true -->
```python
""" THIS CODE CELL IS INTENTIONALLY LEFT EMPTY.  Remove the 'raise NotImplementedError' line. You should leave this cell empty (it is used for autograding). (Do not change the cell type from code to markdown.)"""
# THIS CODE CELL IS INTENTIONALLY LEFT EMPTY --- DO NOT MODIFY THIS CELL
# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" THIS CODE CELL IS INTENTIONALLY LEFT EMPTY.  Remove the 'raise NotImplementedError' line. You should leave this cell empty (it is used for autograding). (Do not change the cell type from code to markdown.)"
```

### [Task 3c] (5 points) Out of the three models, which would you use and why?

<!-- grade_id: task3c_manual_answer, points: 5, grade: true, solution: true -->
```python
""" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"""
# 
## Answer: 
# Based on the results reported, I would choose the Ridge Regression model without polynomial features
# or added regularization. This is because each value r^2, MSE, and MAE are consistent. The
# training values are all slightly higher than the validation, but the validation is still on par meaning
# good generalization by the model. Also, the 3 most important features seem more relevant to one another compared 
# to the other models outputs. Plus, the models trained with polynomial features seem a bit overfitted in the 
# results they output. 

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"
```

## [Task 4] (25 points) Trees and Bagging Ensembles!

### Let's do some cleanup and sanity checking.

```python
# let's do some cleanup and discard all the polynomial features stuff.
if var_exists('all_x_pf'):
    del all_x_pf, train_x_pf, train_y_pf, test_x_pf, test_y_pf, val_x_pf, val_y_pf

# sanity check shapes
train_x.shape, train_y.shape, test_x.shape, test_y.shape, val_x.shape, val_y.shape
assert train_x.shape == (13903, 14) and val_x.shape == (1738, 14) and test_x.shape == (1738, 14)
```

### [Task 4a] (5 points) Suppose a sklearn decision tree has n nodes. How many total splits does it contain? (Hint: think of it as a CS question not an ML question. Also think about edge cases (e.g., n=0).)

<!-- grade_id: task4a_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here to answer by implementing the function.
"""
def num_dt_splits(nodes): # should return the total number of splits in a decision tree. Must return an integer value.
    assert nodes >= 0
    # YOUR CODE HERE
    if nodes == 0:
        return 0
    else:
        nodes = nodes - 1
        return nodes
    # raise NotImplementedError()
```

<!-- grade_id: task4a-tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 4a. """
assert var_exists('num_dt_splits')
assert num_dt_splits(0) == 0 and num_dt_splits(1) == 0
for i in range(0, 100):
    nodes = np.random.randint(0, 1e8)
    splits = num_dt_splits(nodes)
    assert splits>=0 and splits == int(splits)

```

### Let's train a decision tree!

### [Task 4b] (5 points) Train and evaluate two (regression) decision trees with sklearn's DecisionTreeRegressor. The first 'dtmodel' should be trained with default parameters (set the seed). The second 'dtregmodel' should be trained with the default parameters (set the seed) but a max depth of 8. Use r2_mse_mae_eval to print performance of both models.

<!-- grade_id: task4b-code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~5 lines)
"""
from sklearn.tree import DecisionTreeRegressor
# YOUR CODE HERE
dtmodel = DecisionTreeRegressor(random_state=seed)
dtmodel.fit(train_x, train_y)
_ = r2_mse_mae_eval(dtmodel, train_x, train_y, val_x, val_y, pref='[DecisionTreeRegressor]')

dtregmodel = DecisionTreeRegressor(max_depth=8, random_state=seed)
dtregmodel.fit(train_x, train_y)
_ = r2_mse_mae_eval(dtregmodel, train_x, train_y, val_x, val_y, pref='[DecisionRegTreeRegressor]')
# raise NotImplementedError()
```

Output:
```
[DecisionTreeRegressor]Train R^2: 1.000, Val  R^2: 0.982
[DecisionTreeRegressor]Train MSE: 0.000, Val MSE: 545.897
[DecisionTreeRegressor]Train MAE: 0.000, Val MAE: 8.475
[DecisionRegTreeRegressor]Train R^2: 0.963, Val  R^2: 0.961
[DecisionRegTreeRegressor]Train MSE: 1235.614, Val MSE: 1163.905
[DecisionRegTreeRegressor]Train MAE: 15.622, Val MAE: 16.002
```

<!-- grade_id: task4b-checks, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 4b. """
assert var_exists('dtmodel') and var_exists('dtregmodel') 
assert dtmodel.tree_.max_depth >= 8 and dtregmodel.tree_.max_depth == 8

```

### [Task 4c] (5 points) Is the regularized model overfitted? Is it better than the models trained in Tasks 2 and 3? (A few sentences suffice.)

<!-- grade_id: task4c_manual_answer, points: 5, grade: true, solution: true -->
```python
""" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"""
# 
## Answer: 
# The regularized model is not overfitting as much as the original model. This is proven as 
# all of the metrics are around the same values with training being slightly higher except for
# in the case of MAE. Training with deeper tree depths could fix this slight issue. The regularized
# model is better than the models trained in Tasks 2 and 3.This is evident through MSE and MAE being lower
# than each of the other models being compared while also still having comparable, if not higher, 
# r^2 score.

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"
```

### Let's explore bagging.

### [Task 4d] (5 points) Train two ensemble models using sklearn's BaggingRegressor. Use default parameters but be sure to set the seed and set max_samples=0.5. The first ensemble 'lr_bagging' must use 'LinearRegression' as weak learners, whereas the second 'dtr_bagging' must use 'DecisionTreeRegressor' as weak learners.

<!-- grade_id: task4d_code, points: 0, grade: false, solution: true -->
```python
"""Put your code here (~3 lines)
"""
from sklearn.ensemble import BaggingRegressor
# YOUR CODE HERE
lr_bagging = BaggingRegressor(estimator=LinearRegression(), random_state=seed, max_samples=0.5)
lr_bagging.fit(train_x, train_y)

dtr_bagging = BaggingRegressor(estimator=DecisionTreeRegressor(), random_state=seed, max_samples=0.5)
dtr_bagging.fit(train_x, train_y)
# raise NotImplementedError()

_ = r2_mse_mae_eval(lr_bagging, train_x, train_y, val_x, val_y, pref='[Linear Regression Bagging] ')
print()
_ = r2_mse_mae_eval(dtr_bagging, train_x, train_y, val_x, val_y, pref='[DT Regression Bagging aka Random Forest] ')
```

Output:
```
[Linear Regression Bagging] Train R^2: -0.957, Val  R^2: 0.891
[Linear Regression Bagging] Train MSE: 64982.388, Val MSE: 3238.971
[Linear Regression Bagging] Train MAE: 36.519, Val MAE: 32.662

[DT Regression Bagging aka Random Forest] Train R^2: 0.990, Val  R^2: 0.984
[DT Regression Bagging aka Random Forest] Train MSE: 334.994, Val MSE: 478.150
[DT Regression Bagging aka Random Forest] Train MAE: 5.967, Val MAE: 7.918
```

<!-- grade_id: task4d_tests, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 4d. """
assert var_exists('lr_bagging') and var_exists('dtr_bagging') 
assert lr_bagging.random_state == dtr_bagging.random_state and dtr_bagging.random_state == seed

```

### [Task 4e] (5 points) Answer the following questions: 
### 1. Is the linear regression bagging ensemble a better model than linear regression (Task 2)? Is this expected? (We are looking for an explanation of why this approach performs the way it does.)
### 2. How does the random forest compare to all other models in this and previous task?

<!-- grade_id: task4e_manual_answer, points: 5, grade: true, solution: true -->
```python
""" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"""
# 
## Answer: 
# (1) In terms of validation, the linear regression bagging ensemble model generalizes better than the 
# model from Task 2. As we can see, variance is reduced here based on the validation results, but bias
# is high based on the training results. Bagging the models also helps avoid overfitting in the final model
# and lower complexity yields lower variance which leads to better generalization.

# (2) The random forest outperforms all of the previous models in regards to validation results. It proves
# to have better performance overall in accuracy and generalization.

# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
" Remove the 'raise NotImplementedError' line. Write your answer as a comment in the place provided.  (Do not change the cell type from code to markdown.)"
```

---
# [Task 5] \<*For CAI6108MLE Only*\> (25 points) Stacking & MoE. For this task you will implement a (kind of) Stacking/MoE approach. 
# Roughly speaking the approach will be to randomly choose weak learners out of a pool of candidates (determined by the *last* digit of your UFID) and train each one of a randomly selected subset of features. With the weak learners trained, we will train a gating model to learn weights to combine the weak learners' predictions into a single final prediction.

### [Task 5] The following code is provided, you should not modify it. But you should *read* it.

```python
from sklearn.svm import SVR
from sklearn.linear_model import LogisticRegression, Ridge, LassoLars
from sklearn.tree import DecisionTreeRegressor
from sklearn.neighbors import KNeighborsRegressor
from sklearn.gaussian_process import GaussianProcessRegressor

""" Returns candidate weak learners based on UFID digit.
"""
def weak_learner_candidates(ufid_digit):
    if ufid_digit in [1, 3, 6]:
        return [SVR(kernel='poly'), Ridge(), DecisionTreeRegressor(max_depth=10)]
    elif ufid_digit in [2, 4, 7]:
        return [SVR(kernel='poly'), LassoLars(), KNeighborsRegressor(n_neighbors=3)]
    else:
        assert ufid_digit in [0, 5, 8, 9] 
        return [DecisionTreeRegressor(max_depth=8), Ridge(), KNeighborsRegressor(n_neighbors=1)]

""" Randomly sample and return a weak learner out of the candidate pool.
"""
def random_weak_learner(candidates_fn):
    assert candidates_fn is not None

    candidates = candidates_fn()
    assert len(candidates) == 3
    ridx = np.random.randint(0, len(candidates))

    return candidates[ridx]

""" Predicts the value(s) on 'x' given the weak_learners and gating function. (Can you figure out how it works??)
"""
def predict(weak_learners, gating, x, epsilon=1e-4):
    k = len(weak_learners)
    
    gating_weights = np.ones((x.shape[0], k)) * epsilon
    gw = gating.predict_proba(x)
    for i, cidx in enumerate(gating.classes_):
        gating_weights[:,cidx] = gw[:,i]
    gating_weights = gating_weights / np.sum(gating_weights, axis=-1).reshape((-1,1))
    
    pred_y = np.zeros((x.shape[0],))
    for i, tup in enumerate(weak_learners):
        weak_learner, sel_features = tup
        predi = weak_learner.predict(x[:,sel_features])
        pred_y += gating_weights[:,i] * predi
    return pred_y
```

### [Task 5a] (10 points) Fill in the code of 'train_ensemble' and 'train_gating' following the provided instructions. Note: instructions do not spell out what every line of code you write should do --- that is for you to figure out.

<!-- grade_id: task5a-code, points: 0, grade: false, solution: true -->
```python
""" Trains the ensemble of 'num_weak_learners' of randomly chosen weak learners, each selecting a random subset of 'num_sel_features' features.
    Returns a list of tuples composed of the trained weak learner ('weak_learner') and the chosen features ('fidx).
"""
def train_ensemble(tr_x, tr_y, candidates_fn, num_weak_learners=16, num_sel_features=7):
    ret = []
    num_features = train_x.shape[1]
    for i in range(0, num_weak_learners):

        """Put your code here (~3-4 lines)
        """
        # YOUR CODE HERE
        weak_learner = random_weak_learner(candidates_fn)
        fidx = np.random.choice(num_features, num_sel_features, replace=False)
        weak_learner.fit(tr_x[:, fidx], tr_y)
        # raise NotImplementedError()
        
        weak_learner_tup = (weak_learner, fidx)
        ret.append(weak_learner_tup)
    return ret


""" Trains the gating function (logistic regression) using the provided list of weak learners tuples.
    Returns the gating model, errors matrix and best weak learner idx.
"""
def train_gating(weak_learners, tr_x, tr_y):
    k = len(weak_learners)
    errors = np.ones((tr_x.shape[0], k)) * np.inf
    """Put your code here (~3-4 lines). For each weak learner measure its *absolute errors* on tr_{x,y} (put that in 'errors').
    """
    # YOUR CODE HERE
    for i, (weak_learner, sel_features) in enumerate(weak_learners):
        pred_y = weak_learner.predict(tr_x[:, sel_features])
        errors[:, i] = np.abs(pred_y - tr_y)
    # raise NotImplementedError()

    bestidx = np.argmin(errors, axis=-1) # compute the weak learning idx with the lowest error
    
    # train a logistic regression model for the gating
    class_weight = {i:1.0 for i in range(0, k)}
    gating = LogisticRegression(penalty='l2', C=1.0, random_state=seed, max_iter=10000, class_weight=class_weight, multi_class='multinomial')
    gating.fit(tr_x, bestidx)

    return gating, errors, bestidx
```

<!-- grade_id: task5a-checks, points: 10, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 5a. """
assert var_exists('weak_learner_candidates') and var_exists('random_weak_learner') 
assert var_exists('train_gating') and var_exists('train_ensemble') 

```

### [Task 5b] (10 points) Fill in the provided code. Use the functions you previously implemented to train the ensemble and the gating models and make predictions on the validation data.

<!-- grade_id: task5b-code, points: 0, grade: false, solution: true -->
```python
st = time.time()

"""Put your code here (~3-4 lines). 
    1. Split the training data into s1_{x,y} and s2_{x,y} such that 70% ends up in s1 (and 30% in s2).
    2. train the ensemble on s1_{x,y} [do not forget to use the *last* digit of your UFID].
"""
from sklearn.metrics import r2_score
from sklearn.metrics import mean_squared_error
from sklearn.metrics import mean_absolute_error
# YOUR CODE HERE
s1_x, s2_x, s1_y, s2_y = train_test_split(train_x, train_y, test_size=0.3, random_state=seed)
wls = train_ensemble(s1_x, s1_y, lambda: weak_learner_candidates(6), num_weak_learners=16, num_sel_features=7)
# raise NotImplementedError()

# train the gating on s2_{x,y}
gating, errors, bestidx = train_gating(wls, s2_x, s2_y)
# use the predict function to make predictions on the validation data
pred_y = predict(wls, gating, val_x)

"""Put your code here (~3-4 lines). 
    compute the R^2, MSE, MAE on the validation data (using 'pred_y') and store the result as 'val_r2', 'val_mse', 'val_mae'
"""
# YOUR CODE HERE
val_r2 = r2_score(val_y, pred_y)
val_mse = mean_squared_error(val_y, pred_y)
val_mae = mean_absolute_error(val_y, pred_y)
# raise NotImplementedError()

et = time.time()
elapsed = et - st
print('[Val] R^2: {:.3f}, MSE: {:.1f}, MAE: {:.1f} [Elapsed: {:.1f} seconds]'.format(val_r2, val_mse, val_mae, elapsed))
```

Output:
```
[Val] R^2: 0.944, MSE: 1660.4, MAE: 23.1 [Elapsed: 49.1 seconds]
```

<!-- grade_id: task5b_tests, points: 10, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 5b. """

assert var_exists('gating')
assert elapsed < 300.0
assert np.abs(val_r2 - 0.88) <= 0.12 and np.abs(val_mae - 30) <= 30

```

### [Task 5c] (5 points) Now observe that 'num_weak_learners' and 'num_sel_features' in train_ensemble can be considered hyperparameters. Write some code to tune them through a simple grid search and then print the best value for them and the performance of the resulting ensemble model. (For this question best means the one with lowest val R^2.) Is it better than previous models in this homework?

<!-- grade_id: task5c-code, points: 0, grade: false, solution: true -->
```python
"""Put your code here. Call the resulting best hyperparameters 'best_num_weak_learners' and 'best_num_sel_features'
    Store the performance R^2, MSE, MAE as 'val_r2', 'val_mse', 'val_mae'
"""
# YOUR CODE HERE
best_val_r2 = -float('inf')
best_num_weak_learners = None
best_num_sel_features = None

nwlv = [2, 4, 8, 16, 32, 64]
nsfv = [3, 5, 7, 9]

for nwl in nwlv:
    for nsf in nsfv:
        wls = train_ensemble(s1_x, s1_y, lambda: weak_learner_candidates(6), num_weak_learners=nwl, num_sel_features=nsf)
        gating, errors, bestidx = train_gating(wls, s2_x, s2_y)
        pred_y = predict(wls, gating, val_x)
        
        val_r2 = r2_score(val_y, pred_y)
        
        if val_r2 > best_val_r2:
            best_val_r2 = val_r2
            best_num_weak_learners = nwl
            best_num_sel_features = nsf
# raise NotImplementedError()
print('[Val] R^2: {:.3f}, MSE: {:.1f}, MAE: {:.1f} [Best hyperparams: {}, {}]'.format(val_r2, val_mse, val_mae, best_num_weak_learners, best_num_sel_features))
```

Output:
```
[Val] R^2: 0.946, MSE: 1660.4, MAE: 23.1 [Best hyperparams: 8, 7]
```

<!-- grade_id: task5c-checks, points: 5, grade: true, solution: false -->
```python
""" [ASSERTS] Check task 5c. """

assert var_exists('best_num_weak_learners') and var_exists('best_num_sel_features')
assert np.abs(val_r2 - 0.92) <= 0.08 and np.abs(val_mae - 24) <= 24

```